# Assignment 2
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```


<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
import pandas as pd
def proportion_of_education():
    
    # Reading CSV file using pandas
    df = pd.read_csv('assets/NISPUF17.csv')
    
    # Sepearating or spliting the 'EDUC1' from original dataframe & save it in new dataframe...
    df_edu = df['EDUC1']
    
    # This is the total count that we need
    total_row = df_edu.shape[0]
    
    # 1 = less than 12 years
    # 2 = 12 years
    # 3 = grater than 12 years, not a college degree
    # 4 = college graduate

    
    # Let's create the values for every number (4,3,2,1)
    less_than_12years = df_edu.value_counts()[1]
    years12 = df_edu.value_counts()[2]
    great_12years_not_college = df_edu.value_counts()[3]
    college_graduate =  df_edu.value_counts()[4]
    
    # Making the variable for print function.
    proportional = {"less than high school":less_than_12years/total_row,
                    "high school":years12/total_row,
                    "more than high school but not college":great_12years_not_college / total_row,
                    "college":college_graduate/total_row}
    
    # returning the proportional value
    return proportional

    # If we don;t implement any calculation it will raise the implementation error
    raise NotImplementedError()
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    # Reading CSV using pandas

    # This will only work on Coursera Assignment Platfrom
    df = pd.read_csv('assets/NISPUF17.csv')

    # For Google Collab this will work fine
    # df = pd.read_csv('/content/NISPUF17.csv')

    # Creating a dataframe that contain only child who were breast feeded
    milk_feeded = df[df['CBF_01'] == 1]

    # Creating a dataframe that contain only child who were not breast feeded
    milk_not_feeded = df[df['CBF_01'] == 2]

    # Return the tuple value with mean, for both.
    return (milk_feeded['P_NUMFLU'].mean(), milk_not_feeded['P_NUMFLU'].mean())

    raise NotImplementedError()
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():

    # Reading CSV using pandas

    # This will only work on Coursera Assignment Platfrom
    df = pd.read_csv('assets/NISPUF17.csv')

    # For Google Collab this will work fine
    # df = pd.read_csv('/content/NISPUF17.csv')

    # Creating the df that contain one representing sex (SEX), had chicken pox or not (HAD_CPOX), number of varicella dose (P_NUMVRC).
    required_df = df[['P_NUMVRC', 'SEX', 'HAD_CPOX']]

    # Dropping the NaN values from the dataframe in every columns
    droped_nan_values = required_df.dropna() # droped nan  value rows

    # Only Select the vlaues that varicella doeses that are greater than 0
    varicella_doses_greater_0 = droped_nan_values[droped_nan_values['P_NUMVRC'] > 0]

    # Separate the male and female column
    male_df = varicella_doses_greater_0[varicella_doses_greater_0['SEX'] == 1]
    female_df = varicella_doses_greater_0[varicella_doses_greater_0['SEX'] == 2]

    # Separating the male dataframe who had pox and who hadn't pox
    male_had_pox = male_df['HAD_CPOX'].value_counts()[1]
    male_hadnt_pox = male_df.HAD_CPOX.value_counts()[2]

    # Separating the female dataframe who had pox and who hadn't pox
    female_had_pox = female_df['HAD_CPOX'].value_counts()[1]
    female_hadnt_pox = female_df['HAD_CPOX'].value_counts()[2]

    # Find ratio for number of Males/Females who were vaccinated but had chickenpox divided by 
    #   number of Males/Females who were were vaccinated and never had chicken pox.
    return {'male':male_had_pox/male_hadnt_pox,
        'female':female_had_pox / female_hadnt_pox}
    raise NotImplementedError()
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    # Importing all libraries
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # Reading CSV File

    # This will only work on coursera assignment 
    df = pd.read_csv('assets/NISPUF17.csv')
    
    # For Google Collab this will work fine
    # df = pd.read_csv('/content/NISPUF17.csv')

    # Taking the required dataframe
    df = df[['HAD_CPOX', 'P_NUMVRC']]

    df = df[df['HAD_CPOX'] <= 3]
    df = df[df['P_NUMVRC'].notna()]
    corr, pval=stats.pearsonr(df["HAD_CPOX"],df["P_NUMVRC"])

    # returning corr
    return corr
    raise NotImplementedError()
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```